<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Skriptsprachen
### Winterersemester 2023/24
Prof. Dr. Heiner Giefers

### Aufgabe: Die Klasse Matrix

In dieser Aufgabe geht es darum, eine Klasse `Matrix` zur Speicherung und Verarbeitung mathematischer Matrizen zu implementieren. Python besitzt zwar mehrere sequentielle Datentypen (z.B. *tuple*) mit denen Matrizen dargestellt werden können. Allerdings sind sequentielle Datentypen *Container* für allgemeine Objekte und damit nicht sonderlich gut geeignet für mathematische Operationen.

Nehmen wir z.B. zwei Tupel von Tupeln `a` und `b`, die jeweils eine $2\times{}2$ Matrix aus Integer-Werten enthalten. Eine `+` Operation führt nicht wie gewünscht eine Matrix-Addition aus, sondern konkateniert  `a` und `b` zu einem neuen Tupel:

```python
a = ((1,2),(3,4))
b = ((5,5),(5,5))
print(a+b)
```

Ziel dieser Aufgabe ist es, eine Klasse `Matrix` zu programmieren, mit der diese `+`-Operation und weitere Operationen sinnvoll durchgeführt werden kann. Dazu sind eine Reihe von Teilaufgaben nützlich, die Sie in beliebiger Reihenfolge bearbeiten können:

**1. Implementieren Sie einen Konstruktor**

- Intern sollen die Werte der Matrix zeilenweise als eindimensionales Tupel abgelegt werden.
- Die Form (Anzahl der Zeilen und Spalten) der Matrix soll als Attribut gespeichert werden. Es ist sinnvoll, die Form als Tupel anzulegen. Im 2-dimensionalen Fall also als Tupel $(x,y)$ wobei $x$ die Anzahl der Zeilen und $y$ die Anzahl der Spalten ist. $(1,8)$ wäre demnach die Dimension für einen Zeilenvektor mit 8 Elementen, $(8,8)$ die Dimension für eine quadratische Matrix mit 8 Zeilen und Spalten.
- Der Konstruktor soll eine ein- oder zweidimensionale Liste (oder Tupel) als Eingabe bekommen. Wird eine eindimensionale Datenstruktur übergeben, legt der Konstrukter eine Matrix mit einer Zeile an, wird eine zweidimensionale Datenstruktur übergeben, übernimmt der Konstruktor die Dimension der Daten.

Beispiel:
```python
m1 = Matrix([[1,2,3],[4,5,6]])
````
$$
\text{wird zu } m1 = \begin{pmatrix}
1 & 2 & 3 \\
4 & 5 & 6
\end{pmatrix}
$$


**2. Schreiben Sie eine Funktion `__str__` mit der Sie Matrizen formatiert ausgeben können**

- Die Formatierung ist beliebig, aber es sollte klar sein, welche Form die Matrix hat. Sie können sich gerne an der [Matrix Formatierung in Matlab](https://ch.mathworks.com/help/matlab/learn_matlab/matrices-and-arrays.html) orientieren.

**3. Überschreiben Sie den `+`-Operator für die Klasse `Matrix`, sodass das Ergebnis ein neues Matrix-Objekt mit der Summe der beiden Operanden ist**

- Implementieren Sie ähnliche Methode für `-` und `*`
- Bei der Multiplikation mit dem `*`-Zeichen können Sie das elementweise Produkt (Hadamard-Produkt) annehmen. Dabei wird der Wert $a_{i,j}$ der linken Matrix $A$ mit dem Wert $b_{i,j}$ der rechten Matrix $B$ multipliziert. Das Ergebnis bildet den Wert $c_{i,j}$ in der Ergebnismatrix $C$.
- Sie können versuchen, eine *generische Methode* für binäre Operatoren zu schreiben. Damit vermeiden Sie redundanten Code. Eine Möglichkeit dies zu tun ist, das Modul [operator](https://docs.python.org/3/library/operator.html) aus der Standardbibliothek einzusetzen.

**4. Implementieren Sie eine Methode `transpose` mit der Sie die Matrix transponieren können**

- Beim Transponieren vertauschen Sie Zeilen und Spalten
- `transpose` soll ein neues Matrix-Objekt zurückgeben
- Führen Sie eine Kurzschreibweise `A.T` ein mit der Sie die Matrix $A$ Transponieren können. Das *property* Attribut `T` soll die `transpose`-Funktion aufgerufen werden.

**5. Ergänzen Sie die Klasse `Matrix` um eine Methode, mit der die [Matrizenmultiplikation](https://de.wikipedia.org/wiki/Matrizenmultiplikation) zweier Matrizen gebildet werden kann**

- Als Operator können Sie das `@`-Zeichen Verwenden. Die zugehörige *Magic Method* ist `__matmul__(self,x)`

<!-- grade_id: cell-382348fc9bdc8144, points: 0, grade: false, solution: true -->
```python
import operator

class Matrix():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-68944433517f1999, points: 0, grade: true, solution: false -->
```python
m1 = Matrix([[1,2,3],[4,5,6]])
assert (m1.T).data==(1,4,2,5,3,6)
assert (m1 + m1).data==(2,4,6,8,10,12)
assert (m1 - m1).data==(0,0,0,0,0,0)
assert (m1 * m1).data==(1,4,9,16,25,36)
assert (m1 @ m1.T).data==(14.0,32.0,32.0,77.0)
```

```python
m1 = Matrix([[1,2,3],[4,5,6]])
print("m1     :", m1)
print("m1'    :", m1.T)
print("m1 + m1:", m1 + m1)
print("m1 - m1:", m1 - m1)
print("m1 * m1:", m1 * m1)
print("m1 @ m1':", m1 @ m1.T)
```

```python

```